# Student Grades Regression Model

In this notebook you will build a regression model to predict student grades.

## Imports

```python
import pandas as pd
```

## 1. Dataset

This dataset comes from Kaggle and has information about student grades and alcohol usage along with information about their family:

https://www.kaggle.com/uciml/student-alcohol-consumption/kernels

```python
raw_data = pd.read_csv('/data/student-alcohol-consumption/student-mat.csv')
```

```python
raw_data.head()
```

Output:
```
school sex  age address famsize Pstatus  Medu  Fedu     Mjob      Fjob ...  \
0     GP   F   18       U     GT3       A     4     4  at_home   teacher ...   
1     GP   F   17       U     GT3       T     1     1  at_home     other ...   
2     GP   F   15       U     LE3       T     1     1  at_home     other ...   
3     GP   F   15       U     GT3       T     4     2   health  services ...   
4     GP   F   16       U     GT3       T     3     3    other     other ...   

  famrel freetime  goout  Dalc  Walc health absences  G1  G2  G3  
0      4        3      4     1     1      3        6   5   6   6  
1      5        3      3     1     1      3        4   5   5   6  
2      4        3      2     2     3      3       10   7   8  10  
3      3        2      2     1     1      5        2  15  14  15  
4      4        3      2     1     2      5        4   6  10  10  

[5 rows x 33 columns]
```

```python
raw_data.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 395 entries, 0 to 394
Data columns (total 33 columns):
school        395 non-null object
sex           395 non-null object
age           395 non-null int64
address       395 non-null object
famsize       395 non-null object
Pstatus       395 non-null object
Medu          395 non-null int64
Fedu          395 non-null int64
Mjob          395 non-null object
Fjob          395 non-null object
reason        395 non-null object
guardian      395 non-null object
traveltime    395 non-null int64
studytime     395 non-null int64
failures      395 non-null int64
schoolsup     395 non-null object
famsup        395 non-null object
paid          395 non-null object
activities    395 non-null object
nursery       395 non-null object
higher        395 non-null object
internet      395 non-null object
romantic      395 non-null object
famrel        395 non-null int64
freetime      395 non-null int64
goout         395 non-null int64
Dalc          395 no
```

```python
raw_data.columns
```

Output:
```
Index(['school', 'sex', 'age', 'address', 'famsize', 'Pstatus', 'Medu', 'Fedu',
       'Mjob', 'Fjob', 'reason', 'guardian', 'traveltime', 'studytime',
       'failures', 'schoolsup', 'famsup', 'paid', 'activities', 'nursery',
       'higher', 'internet', 'romantic', 'famrel', 'freetime', 'goout', 'Dalc',
       'Walc', 'health', 'absences', 'G1', 'G2', 'G3'],
      dtype='object')
```

## 2. Features

Create a feature `DataFrame`, `X` with the following columns:

* `Dalc` (weekday alcohol consumption)
* `Walc` (weekend alcohol consumption)
* `Medu` (mother's education level)
* `Fedu` (father's education level)
* `traveltime`
* `studytime`
* `goout`
* `romantic` (one hot encoded, with `get_dummies` and `drop_first=True, prefix='romantic'`)
* `higher` (one hot encoded, with `get_dummies` and `drop_first=True, prefix='higher'`)
* `sex` (one hot encoded, with `get_dummies` and `drop_first=True, prefix='higher'`)

<!-- grade_id: grades-regression-d, points: 0, grade: false, solution: true -->
```python
X = raw_data[['Dalc', 'Walc', 'Medu', 'Fedu', 'traveltime', 'studytime', 'goout']]
X = pd.merge(X, pd.get_dummies(raw_data['higher'], drop_first=True, prefix='higher'), left_index=True, right_index=True)
X = pd.merge(X, pd.get_dummies(raw_data['sex'], drop_first=True, prefix='sex'), left_index=True, right_index=True)
X = pd.merge(X, pd.get_dummies(raw_data['romantic'], drop_first=True, prefix='romantic'), left_index=True, right_index=True)
```

```python
X.columns
```

Output:
```
Index(['Dalc', 'Walc', 'Medu', 'Fedu', 'traveltime', 'studytime', 'goout',
       'higher_yes', 'sex_M', 'romantic_yes'],
      dtype='object')
```

<!-- grade_id: grades-regression-e, points: 2, grade: true, solution: false -->
```python
assert list(X.columns)==['Dalc', 'Walc', 'Medu', 'Fedu', 'traveltime', 'studytime', 'goout',
       'higher_yes', 'sex_M', 'romantic_yes']
```

```python
X.head()
```

Output:
```
Dalc  Walc  Medu  Fedu  traveltime  studytime  goout  higher_yes  sex_M  \
0     1     1     4     4           2          2      4           1      0   
1     1     1     1     1           1          2      3           1      0   
2     2     3     1     1           1          2      2           1      0   
3     1     1     4     2           1          3      2           1      0   
4     1     2     3     3           1          2      2           1      0   

   romantic_yes  
0             0  
1             0  
2             0  
3             1  
4             0
```

Create the target column `y` from the `G3` column (total grade):

<!-- grade_id: grades-regression-g, points: 0, grade: false, solution: true -->
```python
y = raw_data['G3']
```

<!-- grade_id: grades-regression-h, points: 2, grade: true, solution: false -->
```python
assert list(y.value_counts().values)==[56, 47, 38, 33, 32, 31, 31, 28, 27, 16, 15, 12,  9,  7,  6,  5,  1,
        1]
```

```python
from sklearn.model_selection import train_test_split
```

<!-- grade_id: grades-regression-j, points: 0, grade: false, solution: true -->
```python
Xtrain, Xtest, ytrain, ytest = train_test_split(X, y, test_size=.3)
```

<!-- grade_id: grades-regression-k, points: 2, grade: true, solution: false -->
```python
assert Xtrain.shape==(276,10)
assert Xtest.shape==(119,10)
assert ytrain.shape==(276,)
assert ytest.shape==(119,)
```

## Regression model

```python
from sklearn.linear_model import LinearRegression, Lasso, Ridge
from sklearn.ensemble import RandomForestRegressor
from sklearn.preprocessing import PolynomialFeatures
from sklearn.pipeline import make_pipeline
from sklearn.metrics import r2_score
```

In the following cells create and tune regression models using the following models:

* `LinearRegression`
* `RandomForestRegression`
* `Lasso`

For each of the models:

* Create a pipeline with a `PolynomialFeatures` preprocessor first and model second.
* Compute the $R^2$ score for both the training and test datasets.
* Tune model parameters, including the polynomial degree to balance the bias and variance of the model.

Create, fit, tune and predict using the `LinearRegression` model here:

<!-- grade_id: grades-regression-m, points: 2, grade: true, solution: true -->
```python
model_lr = make_pipeline(PolynomialFeatures(1), LinearRegression())
model_lr.fit(Xtrain, ytrain)
ypred = model_lr.predict(Xtest)
```

Compute and print the training and test $R^2$ score here: 

<!-- grade_id: grades-regression-n, points: 2, grade: true, solution: true -->
```python
r2_lr = r2_score(ytest, ypred)
print(r2_lr)
```

Output:
```
0.0401610155764
```

Create, fit, tune and predict using the `RandomForestRegression` model here:

<!-- grade_id: grades-regression-o, points: 2, grade: true, solution: true -->
```python
model_rf = make_pipeline(PolynomialFeatures(5), RandomForestRegressor())
model_rf.fit(Xtrain, ytrain)
ypred = model_rf.predict(Xtest)
```

Compute and print the training and test $R^2$ score here: 

<!-- grade_id: grades-regression-p, points: 2, grade: true, solution: true -->
```python
r2_rf = r2_score(ytest, ypred)
print(r2_rf)
```

Output:
```
-0.250672055075
```

Create, fit, tune and predict using the `Ridge` model here:

<!-- grade_id: grades-regression-q, points: 2, grade: true, solution: true -->
```python
model_ridge = make_pipeline(PolynomialFeatures(2), Ridge(alpha=.2))
model_ridge.fit(Xtrain, ytrain)
ypred = model_ridge.predict(Xtest)
```

Compute and print the training and test $R^2$ score here: 

<!-- grade_id: grades-regression-r, points: 2, grade: true, solution: true -->
```python
r2_ridge = r2_score(ytest, ypred)
print(r2_ridge)
```

Output:
```
-0.373844967079
```

```python

```