# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: projecteuler52a, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    str_x = str(x)
    str_y = str(y)
    if len(str_x) != len(str_y):
        return False
    for i in str_x:
        if i not in str_y:
            return False
    return True
```

<!-- grade_id: projecteuler52b, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: projecteuler52c, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
num = 100000 #has to have 6 digits
while(True):
    if same_digits(num, 2*num):
        #print("2x: ", 2*num)
        if same_digits(num, 3*num):
            #print("3x: ", 3*num)
            if same_digits(num, 4*num):
                #print("4x: ", 4*num)
                if same_digits(num, 5*num):
                    #print("5x: ", 5*num)
                    if same_digits(num, 6*num):
                        #print("6x: ", 6*num)
                        print("N: ", num)
                        break
    num += 1
```

Output:
```
N:  142857
```